## Week 6 Assignment - W200 Python for Data Science, UC Berkeley MIDS

Write code in this jupyter Notebook to solve each of the following problems. Each problem should have its solution in a separate cell. Please upload this **Notebook**, your **scrabble.py** file, the **sowpods.txt** file and your **score_word** module with your solutions to your GitHub repository in your SUBMISSIONS/week_06 folder by 11:59PM PST the night before class.

## Objectives:

- Read and understand PEP 8 standards
- Use all of your previously gained knowledge together on a single program
- Demonstrate how to import a user made module and function into python from another .py file
- Demonstrate how to input command line arguments into a .py file

## 6-1. PEP 8 Style Guide (reading and response)

Your first task for this week is to write a **250** word reading response to the article below. In addition, please list **3 questions** that you have from the article. Please write your response in a markdown cell, below this cell in the notebook.

The writing response is a free response, so you may write about your reactions. An interesting thing that you saw in the article, something that really stuck out to you, etc.

**Article**: [The PEP 8 Style Guide](https://www.python.org/dev/peps/pep-0008). This document is really important for Python coders because it describes best practices and customs for how one should write Python code. Please read it all and prepare your questions for class.

<!-- grade_id: cell-eaae902b022d20da, points: 20, grade: true, solution: true -->
### Response

The main finding while reading this document was actually not related to style, but that it is possible to have functions and variables annotations and type checking. While dynamic typing brings a lot of flexibility and may speed up the code writing process, it may be more prone to bugs by not checking, for example, if the output of a function is of the type required to perform other operations. Also, sometimes it is useful for the writer of the code to guide the user on what the expected type is, and for the user to have some documentation of this, or to be reminded when the type is not what the creator of the code intended it to be.

More related to the guide itself, I found very interesting and useful the emphasis that is placed in readability. Specially, I was amazed by the simplicity and brilliance of the insight that code is more often read than written. This should had been obvious from my experience, since I spend much more time reading other people's or my code than writing, and in consequence I use to spend a lot of time making the code more succint, for example, and not as much as I should making it more readable. Therefore, one of my purposes from now on is, besides applying these rules, to have the reader in mind while writing code.

There are also a couple of questions I had while reading the document for which I could not find a very clear/consistent answer online:
- Is there a particular reason to recommend indenting using 4 spaces instead of tab, or is it just an arbitrary convention that was chosen because it is useful that everyone do the same instead of having different uses?
- What is meant by **usage** and **implementation** in the sentence “Names that are visible to the user as public parts of the API should follow conventions that reflect **usage** rather than **implementation**"?
- Lastly, and more as a comment than a question, when writing code, I’m usually in a situation where some programmers I work with speak Spanish but not English or vice versa. In this case, it is not so clear in which language I should write comments and I often find myself coming up with decision rules for which the recommendation to almost always write comments in English is not very useful.


## 6-2. Cheating at Scrabble

Write a Python script that takes a Scrabble rack as a command-line argument and prints all valid Scrabble words that can be constructed from that rack, along with their Scrabble scores, sorted by score. Valid Scrabble words are provided in the data source below. A Scrabble rack is made up of a maximum of any 7 characters.

Below are the requirements for the program:
- This needs to be able to be run as a command line tool as shown below (not an input statement!)
- Please name the python file: `scrabble.py`
- Allow anywhere from 2-7 character tiles to be inputted 
- Do not worry about the number of the same tiles (e.g. a user is allowed to input ZZZZZQQ)
- Output the **total** list of words as (score, word) tuples, sorted by the score as shown below
- Output at the end the 'Total number of words:' that can be made with the letters
- Please include a function called `score_word` in a separate module. Import this function into your main solution code.
- You need to handle input errors from the user and suggest what that error might be caused by and how to fix it (i.e. a helpful error message)
- Implement wildcards as either `*` or `?`. That is, let the user specify a wildcard character that can take any value. There can be a total of two wild cards in any user input (one of each character). Only use the `*` and `?` as wildcard characters
- Wildcard characters are scored as 0 points, just like in the real Scrabble game
- Your program should take less than a minute to run with 2 wildcards in the input - if it is more time than that your algorithm is not optimized very well!
- Write docstrings for the functions and puts comments in your code.

Extra Credit (+10 points):
Allow a user to specify that a certain letter has to be at a certain location. Your program must work without it so this is completely optional. For the extra credit, locations of certain letters must be specified at the command line, it may not be some sort of user prompt.  (Please put a sample of how to run your extra credit & comments in the extra credit cell of this notebook below!)

An example invocation and output:
```
$ python scrabble.py ZAEFIEE
(17, feeze)
(17, feaze)
(16, faze)
(15, fiz)
(15, fez)
(12, zee)
(12, zea)
(11, za)
(6, fie)
(6, fee)
(6, fae)
(5, if)
(5, fe)
(5, fa)
(5, ef)
(2, ee)
(2, ea)
(2, ai)
(2, ae)
Total number of words: 19
```

The Data
http://courses.cms.caltech.edu/cs11/material/advjava/lab1/sowpods.zip contains all words in the official SOWPODS word list, one word per line. You should download the word file and keep it in your repository so that the program is standalone (instead of accessing it over the web from Python).

You can read data from a text file with the following code:

```
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]
```
This will show the first 6 words:
```
print(data[0:6])
```

Please use the dictionary below containing the letters and their Scrabble values:

```
scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}
```

Tips:
- We recommend that you work on this and try to break down the problems into steps on your own before writing any code. Once you've scoped generally what you want to do, start writing some code and if you get stuck, take a step back and go back to thinking about the problem rather than trying to fix lots of errors at the code level. You should only use the Python standard library in this assignment, however any tool in the standard library is fair game.
- If you keep getting stuck check out: https://openhatch.org/wiki/Scrabble_challenge. This is where we got the idea for this assignment and it provides some helpful tips for guiding you along the way. If that link doesn't work you can use the Google cached version here. However, we would recommend that you try to implement this first before looking at the hints on the website.

Good luck!

### The code below will test your command line implementation of the scrabble.py code. We've made some of these tests are available for you to try!

```python
# Code for the testing

import subprocess
from nose.tools import assert_equal 
from nose.tools import assert_true
from nose.tools import assert_greater
from nose.tools import assert_less
```

<!-- grade_id: cell-f55d7e98eba40a10, points: 20, grade: true, solution: false -->
```python
""" Checks that the code runs and checks one user error messages 
    (this is just one of many user errors you should be checking for!)
"""

!python scrabble.py  # no rack error


```

```python
""" Checks that the code runs and checks one user error messages 
    (this is just one of many user errors you should be checking for!)
"""

!python scrabble.py
```

Output:
```
Traceback (most recent call last):
  File "scrabble.py", line 30, in <module>
    You did not enter characters argument.")
Exception: Characters should be entered.        You did not enter characters argument.
```

<!-- grade_id: cell-1893c284e867334a, points: 10, grade: true, solution: false -->
```python
""" Does not fail due to trivial mistakes and takes correct wildcard characters """

!python scrabble.py PENguin    # does not fail due to case
!python scrabble.py PEN*?in    # takes wildcards
```

```python
""" Does not fail due to trivial mistakes and takes correct wildcard characters """

!python scrabble.py PENguin
!python scrabble.py "PEN*?in"
```

Output:
```
(10, penguin)
(9, pening)
(8, genip)
(8, unpeg)
(7, ingenu)
(7, penni)
(7, ping)
(7, pung)
(7, unpen)
(7, unpin)
(6, gip)
(6, gup)
(6, peg)
(6, pein)
(6, peni)
(6, pig)
(6, pine)
(6, pug)
(5, ennui)
(5, genu)
(5, gien)
(5, ginn)
(5, nep)
(5, nip)
(5, pen)
(5, pie)
(5, pin)
(5, piu)
(5, pun)
(4, eng)
(4, gen)
(4, gie)
(4, gin)
(4, gnu)
(4, gue)
(4, gun)
(4, neg)
(4, nine)
(4, pe)
(4, pi)
(4, up)
(3, gi)
(3, gu)
(3, inn)
(3, nie)
(3, nun)
(3, ug)
(3, uni)
(2, en)
(2, in)
(2, ne)
(2, nu)
(2, un)
Total number of words: 53
(7, enprint)
(7, neaping)
(7, ninepin)
(7, opening)
(7, pannier)
(7, pantine)
(7, peaning)
(7, peening)
(7, peining)
(7, pending)
(7, penguin)
(7, pening)
(7, penni)
(7, pennia)
(7, pennied)
(7, pennies)
(7, pennill)
(7, pennine)
(7, penning)
(7, pennis)
(7, pension)
(7, pfennig)
(7, pinbone)
(7, pinene)
(7, pinenes)
(7, pinken)
(7, pinkens)
(7, pinnace)
(7, pinnae)
(7, pinnate)
(7, pinned)
(7, pinner)
(7, pinners)
(7, pinnet)
(7, pinnets)
(7, pinnie)
(7, pinnies)
(7, pin
```

<!-- grade_id: cell-7f1c7fe54a38a19e, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """

#test the code in the command line
!python scrabble.py PENGUIN

```

Output:
```
(10, penguin)
(9, pening)
(8, genip)
(8, unpeg)
(7, ingenu)
(7, penni)
(7, ping)
(7, pung)
(7, unpen)
(7, unpin)
(6, gip)
(6, gup)
(6, peg)
(6, pein)
(6, peni)
(6, pig)
(6, pine)
(6, pug)
(5, ennui)
(5, genu)
(5, gien)
(5, ginn)
(5, nep)
(5, nip)
(5, pen)
(5, pie)
(5, pin)
(5, piu)
(5, pun)
(4, eng)
(4, gen)
(4, gie)
(4, gin)
(4, gnu)
(4, gue)
(4, gun)
(4, neg)
(4, nine)
(4, pe)
(4, pi)
(4, up)
(3, gi)
(3, gu)
(3, inn)
(3, nie)
(3, nun)
(3, ug)
(3, uni)
(2, en)
(2, in)
(2, ne)
(2, nu)
(2, un)
Total number of words: 53
```

<!-- grade_id: cell-bb7a872eda80e6a5, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-8d2de0735f148301, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """
!python scrabble.py PENGU*?

```

```python
""" The code should produce a list of all words from the rack with scores """
!python scrabble.py "PENGU*?"

```

Output:
```
(8, engulph)
(8, expugn)
(8, expugns)
(8, expunge)
(8, penguin)
(8, plunge)
(8, plunged)
(8, plunger)
(8, plunges)
(8, puering)
(8, pungent)
(8, pungle)
(8, pungled)
(8, pungles)
(8, repugn)
(8, repugns)
(8, spueing)
(8, spunge)
(8, spunges)
(8, unpaged)
(8, unpeg)
(8, unpegs)
(8, upgone)
(7, duping)
(7, eggcup)
(7, epigon)
(7, gauped)
(7, gauper)
(7, genip)
(7, genips)
(7, getup)
(7, getups)
(7, gipsen)
(7, gowpen)
(7, guimpe)
(7, gulped)
(7, gulper)
(7, gumped)
(7, hangup)
(7, impugn)
(7, oppugn)
(7, ouping)
(7, panged)
(7, pangen)
(7, peeing)
(7, peenge)
(7, penang)
(7, pengo)
(7, pengos)
(7, pening)
(7, pieing)
(7, pigeon)
(7, pignus)
(7, pignut)
(7, pigpen)
(7, pinged)
(7, pinger)
(7, pingle)
(7, plague)
(7, pleugh)
(7, plonge)
(7, ponged)
(7, pongee)
(7, popgun)
(7, potgun)
(7, progun)
(7, pudge)
(7, pudges)
(7, pugged)
(7, puggie)
(7, puggle)
(7, pugree)
(7, puking)
(7, puling)
(7, pung)
(7, punga)
(7, pungas)
(7, pungs)
(7, purge)
(7, purged)
(7, purger)
(7, purges)
(7, puring)
```

<!-- grade_id: cell-16d4589caf032f51, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-ceab3d4433043319, points: 10, grade: true, solution: false -->
```python
""" The code should run in seconds to a few minutes """
import time
start=time.time()

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])

tot_time=time.time()-start
print('Total time was {} seconds'.format(tot_time))
assert_less(tot_time, 300)
```

Output:
```
Total time was 2.407376766204834 seconds
```

<!-- grade_id: cell-157fccf420d13535, points: 10, grade: true, solution: true -->
```python
"""Implement extra credit call run the code here
   If this cell isnt filled out - we'll assume the extra credit wasn't done
   Please write a comment on how to use your extra credit syntax also!
"""
# YOUR EXTRA CREDIT CODE HERE
```

## If you have feedback for this homework, please submit it using the link below:

http://goo.gl/forms/74yCiQTf6k